# NumPy Workshop: Auto-Graded Assignment
This notebook contains exercises covering key NumPy concepts. Each question includes automated tests.

🔒 **Do not modify test cells.**

**Instructions:** Fill in each `# Your solution here` cell and run the test cell that follows to validate.

```python
import numpy as np
```

### Problem: create_basic_array
Create a NumPy array from [10, 20, 30].

<!-- grade_id: create_basic_array, points: 5, grade: true, solution: true -->
```python
# ID: create_basic_array
# POINTS: 5
arr = 
```

<!-- grade_id: create_basic_array_test, points: 0, grade: true, solution: false -->
```python
# ID: create_basic_array_test
# This cell is autograded.
assert np.array_equal(arr, np.array([10, 20, 30]))
```

### Problem: create_eye_matrix
Create a 4x4 identity matrix and multiply it by 5. Store in `eye_mat`.

<!-- grade_id: create_eye_matrix, points: 5, grade: true, solution: true -->
```python
# ID: create_eye_matrix
# POINTS: 5
eye_mat = 
```

<!-- grade_id: create_eye_matrix_test, points: 0, grade: true, solution: false -->
```python
# ID: create_eye_matrix_test
# This cell is autograded.
assert np.array_equal(eye_mat, np.eye(4)*5)
```

### Problem: linspace_array
Create an array `even_space` with 20 values evenly spaced from 5 to 15.

<!-- grade_id: linspace_array, points: 5, grade: true, solution: true -->
```python
# ID: linspace_array
# POINTS: 5
even_space = 
```

<!-- grade_id: linspace_array_test, points: 0, grade: true, solution: false -->
```python
# ID: linspace_array_test
# This cell is autograded.
assert even_space.shape == (20,) and np.isclose(even_space[0], 5) and np.isclose(even_space[-1], 15)
```

### Problem: zero_shape
Create a 3x3 array of zeros and store in `z`.

<!-- grade_id: zero_shape, points: 5, grade: true, solution: true -->
```python
# ID: zero_shape
# POINTS: 5
z = 
```

<!-- grade_id: zero_shape_test, points: 0, grade: true, solution: false -->
```python
# ID: zero_shape_test
# This cell is autograded.
assert z.shape == (3,3) and z.ndim == 2
```

### Problem: three_dim_array
Create a (2,3,4) array of ones in `arr3d`.

<!-- grade_id: three_dim_array, points: 5, grade: true, solution: true -->
```python
# ID: three_dim_array
# POINTS: 5
arr3d = 
```

<!-- grade_id: three_dim_array_test, points: 0, grade: true, solution: false -->
```python
# ID: three_dim_array_test
# This cell is autograded.
assert arr3d.shape == (2,3,4) and arr3d.ndim == 3
```

### Problem: convert_dtype
Convert float array `[1.1, 2.2, 3.3]` to integers in `int_arr`.

<!-- grade_id: convert_dtype, points: 5, grade: true, solution: true -->
```python
# ID: convert_dtype
# POINTS: 5
int_arr = 
```

<!-- grade_id: convert_dtype_test, points: 0, grade: true, solution: false -->
```python
# ID: convert_dtype_test
# This cell is autograded.
assert np.array_equal(int_arr, np.array([1,2,3])) and int_arr.dtype == int
```

### Problem: reshape_basic
Reshape array from 1D (12,) to 2D (3,4).

<!-- grade_id: reshape_basic, points: 5, grade: true, solution: true -->
```python
# ID: reshape_basic
# POINTS: 5
reshaped = np.arange(12).reshape(3,4)
```

<!-- grade_id: reshape_basic_test, points: 0, grade: true, solution: false -->
```python
# ID: reshape_basic_test
# This cell is autograded.
assert reshaped.shape == (3,4)
```

### Problem: resize_array
Resize (4,3) array to (2,6).

<!-- grade_id: resize_array, points: 5, grade: true, solution: true -->
```python
# ID: resize_array
# POINTS: 5
resized = np.resize(np.arange(12), (2,6))
```

<!-- grade_id: resize_array_test, points: 0, grade: true, solution: false -->
```python
# ID: resize_array_test
# This cell is autograded.
assert resized.shape == (2,6)
```

### Problem: reshape_infer
Use -1 to reshape 2D array with 9 elements into shape (3,3).

<!-- grade_id: reshape_infer, points: 5, grade: true, solution: true -->
```python
# ID: reshape_infer
# POINTS: 5
reshaped_auto = np.arange(9).reshape(-1, 3)
```

<!-- grade_id: reshape_infer_test, points: 0, grade: true, solution: false -->
```python
# ID: reshape_infer_test
# This cell is autograded.
assert reshaped_auto.shape == (3,3)
```

### Problem: index_1d
Extract values from index 2 to 5 (inclusive) of array 0–9.

<!-- grade_id: index_1d, points: 5, grade: true, solution: true -->
```python
# ID: index_1d
# POINTS: 5
slice_1d = np.arange(10)[2:6]
```

<!-- grade_id: index_1d_test, points: 0, grade: true, solution: false -->
```python
# ID: index_1d_test
# This cell is autograded.
assert np.array_equal(slice_1d, np.array([2,3,4,5]))
```

### Problem: boolean_filter
Filter values > 5 from array 0–10 into `filtered`.

<!-- grade_id: boolean_filter, points: 5, grade: true, solution: true -->
```python
# ID: boolean_filter
# POINTS: 5
filtered = np.arange(11)[np.arange(11) > 5]
```

<!-- grade_id: boolean_filter_test, points: 0, grade: true, solution: false -->
```python
# ID: boolean_filter_test
# This cell is autograded.
assert np.array_equal(filtered, np.array([6,7,8,9,10]))
```

### Problem: slice_2d
From a 3x3 array, extract middle row into `middle_row`.

<!-- grade_id: slice_2d, points: 5, grade: true, solution: true -->
```python
# ID: slice_2d
# POINTS: 5
a = np.arange(9).reshape(3,3)
middle_row = a[1,:]
```

<!-- grade_id: slice_2d_test, points: 0, grade: true, solution: false -->
```python
# ID: slice_2d_test
# This cell is autograded.
assert np.array_equal(middle_row, np.array([3,4,5]))
```

### Problem: elementwise_add
Add arrays `[1,2,3]` and `[4,5,6]` into `sum_result`.

<!-- grade_id: elementwise_add, points: 5, grade: true, solution: true -->
```python
# ID: elementwise_add
# POINTS: 5
sum_result = np.array([1,2,3]) + np.array([4,5,6])
```

<!-- grade_id: elementwise_add_test, points: 0, grade: true, solution: false -->
```python
# ID: elementwise_add_test
# This cell is autograded.
assert np.array_equal(sum_result, np.array([5,7,9]))
```

### Problem: dot_product
Calculate dot product of `[1,2]` and `[3,4]`.

<!-- grade_id: dot_product, points: 5, grade: true, solution: true -->
```python
# ID: dot_product
# POINTS: 5
dot_result = np.dot([1,2],[3,4])
```

<!-- grade_id: dot_product_test, points: 0, grade: true, solution: false -->
```python
# ID: dot_product_test
# This cell is autograded.
assert dot_result == 11
```

### Problem: matmul_2d
Use matmul to multiply 2D arrays (2x3) and (3x2).

<!-- grade_id: matmul_2d, points: 5, grade: true, solution: true -->
```python
# ID: matmul_2d
# POINTS: 5
A = np.ones((2,3))
B = np.ones((3,2))
product = np.matmul(A, B)
```

<!-- grade_id: matmul_2d_test, points: 0, grade: true, solution: false -->
```python
# ID: matmul_2d_test
# This cell is autograded.
assert product.shape == (2,2)
```

### Problem: rand_1d
Generate a 1D array of 5 random floats in `rand_floats`.

<!-- grade_id: rand_1d, points: 5, grade: true, solution: true -->
```python
# ID: rand_1d
# POINTS: 5
rand_floats = np.random.rand(5)
```

<!-- grade_id: rand_1d_test, points: 0, grade: true, solution: false -->
```python
# ID: rand_1d_test
# This cell is autograded.
assert rand_floats.shape == (5,)
```

### Problem: randint_2d
Generate 3x3 array of random ints from 0–9.

<!-- grade_id: randint_2d, points: 5, grade: true, solution: true -->
```python
# ID: randint_2d
# POINTS: 5
randint_arr = np.random.randint(0,10,size=(3,3))
```

<!-- grade_id: randint_2d_test, points: 0, grade: true, solution: false -->
```python
# ID: randint_2d_test
# This cell is autograded.
assert randint_arr.shape == (3,3)
```

### Problem: randn_stats
Generate 1000 values from normal distribution and compute mean.

<!-- grade_id: randn_stats, points: 5, grade: true, solution: true -->
```python
# ID: randn_stats
# POINTS: 5
norm_vals = np.random.randn(1000)
norm_mean = np.mean(norm_vals)
```

<!-- grade_id: randn_stats_test, points: 0, grade: true, solution: false -->
```python
# ID: randn_stats_test
# This cell is autograded.
assert abs(norm_mean) < 0.2
```

### Problem: find_min_max
Find min/max of `[3,1,4,2]`.

<!-- grade_id: find_min_max, points: 5, grade: true, solution: true -->
```python
# ID: find_min_max
# POINTS: 5
arr = np.array([3,1,4,2])
min_val = np.min(arr)
max_val = np.max(arr)
```

<!-- grade_id: find_min_max_test, points: 0, grade: true, solution: false -->
```python
# ID: find_min_max_test
# This cell is autograded.
assert min_val == 1 and max_val == 4
```

### Problem: stats_ops
Compute mean, median, std of `[1,2,3,4]`.

<!-- grade_id: stats_ops, points: 5, grade: true, solution: true -->
```python
# ID: stats_ops
# POINTS: 5
data = np.array([1,2,3,4])
data_mean = np.mean(data)
data_median = np.median(data)
data_std = np.std(data)
```

<!-- grade_id: stats_ops_test, points: 0, grade: true, solution: false -->
```python
# ID: stats_ops_test
# This cell is autograded.
assert data_mean == 2.5 and data_median == 2.5
```

### Problem: use_where
Find indices in `[1,2,3,4,5]` where values > 3.

<!-- grade_id: use_where, points: 5, grade: true, solution: true -->
```python
# ID: use_where
# POINTS: 5
indices = np.where(np.array([1,2,3,4,5]) > 3)
```

<!-- grade_id: use_where_test, points: 0, grade: true, solution: false -->
```python
# ID: use_where_test
# This cell is autograded.
assert np.array_equal(indices[0], np.array([3,4]))
```

### Problem: scalar_broadcast
Add scalar 10 to `[1,2,3]`.

<!-- grade_id: scalar_broadcast, points: 5, grade: true, solution: true -->
```python
# ID: scalar_broadcast
# POINTS: 5
broadcasted = np.array([1,2,3]) + 10
```

<!-- grade_id: scalar_broadcast_test, points: 0, grade: true, solution: false -->
```python
# ID: scalar_broadcast_test
# This cell is autograded.
assert np.array_equal(broadcasted, np.array([11,12,13]))
```

### Problem: broadcast_2d
Multiply 1D array `[1,2,3]` with each row of a 2D array.

<!-- grade_id: broadcast_2d, points: 5, grade: true, solution: true -->
```python
# ID: broadcast_2d
# POINTS: 5
A = np.ones((3,3))
B = np.array([1,2,3])
result = A * B
```

<!-- grade_id: broadcast_2d_test, points: 0, grade: true, solution: false -->
```python
# ID: broadcast_2d_test
# This cell is autograded.
assert result.shape == (3,3) and np.array_equal(result[0], np.array([1,2,3]))
```

### Problem: complex_broadcast
Use broadcasting to add column vector to each column of 3x3 matrix.

<!-- grade_id: complex_broadcast, points: 5, grade: true, solution: true -->
```python
# ID: complex_broadcast
# POINTS: 5
mat = np.ones((3,3))
vec = np.array([[1],[2],[3]])
res = mat + vec
```

<!-- grade_id: complex_broadcast_test, points: 0, grade: true, solution: false -->
```python
# ID: complex_broadcast_test
# This cell is autograded.
assert res.shape == (3,3) and np.array_equal(res[:,0], np.array([2,3,4]))
```

### Problem: final_project
Mini-project: Generate random 5x5 array, scale by 2, add 5x5 identity matrix.

<!-- grade_id: final_project, points: 10, grade: true, solution: true -->
```python
# ID: final_project
# POINTS: 10
R = np.random.rand(5,5)
result = 2*R + np.eye(5)
```

<!-- grade_id: final_project_test, points: 0, grade: true, solution: false -->
```python
# ID: final_project_test
# This cell is autograded.
assert result.shape == (5,5)
```